### Evaluating a Polynomial [7 points]

A polynomial in one variable (*univariate*) is of the form `p(x) = c(n) × xⁿ + c(n-1) × xⁿ⁻¹ + ... + c(1) × x + c(0)`, where the `c(i)` are the _coefficients_ and `n` is the _degree_.

Assuming that `N ≥ 0`, the array `c: 0 .. N – 1 → integer` has the coefficients, and `x` is an integer, write a program that, for integer variable `s`, establishes the postcondition

```algorithm
s = (∑ i ∈ 0 .. N - 1 • c(i) × xⁱ)
```

##### 1. With exponentiation [4 points]

The program should contain only one loop and use exponentiation. Give the programs textually and add the appropriate annotation! Use as a template:

```algorithm
{true}
Your annotated algorithm here
{s = (∑ i ∈ 0 .. N – 1 • c(i) × xⁱ)}
```

Give the correctness conditions and prove them!

<!-- grade_id: cell-0ad0faf785be6234, points: 4, grade: true, solution: true -->
```algorithm
{true}
s, j := 0, 0
{P: s = (∑ i ∈ 0 .. j – 1 • c(i) × xⁱ) ∧ 0 ≤ j ≤ N}
while j < N do
	s, j := s + c(j) × xʲ, j + 1
{Q: s = (∑ i ∈ 0 .. N – 1 • c(i) × xⁱ)}
```

The correctness conditions are:
1. `{true} s, j := 0, 0 {P}`
2. `{j < N ∧ P} s, j := s + c(j) × xʲ, j + 1 {P}`
3. `j ≥ N ∧ P ⇒ Q`

The proofs are:

1.
```algorithm
    {true} s, j := 0, 0 {P}
≡       {rule for assignments}
    true ⇒ P[s, j := 0, 0]
≡       {logic, definition of P, substitution}
    0 = (∑ i ∈ 0 .. – 1 • c(i) × xⁱ) ∧ 0 ≤ 0 ≤ N
≡       {empty sum, assumption N ≥ 0}
    true
```

2.  
```algorithm
    {j < N ∧ P} s, j := s + c(j) × xʲ, j + 1 {P}
≡       {rule for assignments}
    j < N ∧ P ⇒ P[s, j := s + c(j) × xʲ, j + 1]
≡       {definition of P, substitution, arithmetic}
    j < N ∧ P ⇒ s + c(j) × xʲ = (∑ i ∈ 0 .. j • c(i) × xⁱ) ∧ 0 ≤ j + 1 ≤ N
≡       {j < N ≡ j + 1 ≤ N, P ⇒ 0 ≤ j + 1, splitting ∑}
    j < N ∧ P ⇒ s + c(j) × xʲ = (∑ i ∈ 0 .. j – 1 • c(i) × xⁱ) + c(j) × xʲ
≡       {definition of P, arithmetic}
    true
```

3.  
```algorithm
    j ≥ N ∧ P ⇒ Q
≡       {rule for assignments}
    j ≥ N ∧ P
≡       {P ⇒ j ≤ N, definition of P}
    j = N ∧ s = (∑ i ∈ 0 .. j – 1 • c(i) × xⁱ)
⇒       {subsituting j for N, definition of Q}
    Q
```

##### 2. Without exponentiation [2 points]

The program should contain only one loop and use only addition and multiplication, no exponentiation. Add appropriate annotation. You do not need to prove that the program is correct, but the loop invariant has to be strong enough to allow a proof (you may want to carry out the proofs for yourself). Use the above template.

<!-- grade_id: cell-afb351faef05f180, points: 2, grade: true, solution: true -->
A new variable, `p`, is introduced with the invariant `xʲ = p`:

```algorithm
{true}
s, p, j := 0, 1, 0
{P: s = (∑ i ∈ 0 .. j – 1 • c(i) × xⁱ) ∧ xʲ = p ∧ 0 ≤ j ≤ N}
while j < N do
	s, p, j := s + c(j) × p, p × x, j + 1
{Q: s = (∑ i ∈ 0 .. N – 1 • c(i) × xⁱ)}
```

##### 3. Implementation [1 point]

Implement your algorithm without exponentions in Python!

<!-- grade_id: cell-293ee825ca9be7d7, points: 0, grade: true, solution: true -->
```python
def poly(c, x):
    ### BEGIN SOLUTION
    s, p = 0, 1
    for j in range(len(c)):
        s, p = s + c[j] * p, p * x
    return s
    ### END SOLUTION
```

<!-- grade_id: cell-a563d7766ede7e8a, points: 1, grade: true, solution: false -->
```python
assert poly([1, 2, 3], 3) == 34
assert poly([7], 6) == 7
assert poly([], 3) == 0
assert poly([2, 0, 0], 7) == 2
assert poly([0, 0, 2], 7) == 98
```